<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 94093  -->
  <question type="stack">
    <name>
      <text>Place the estimated value of mode in stacked dotplot</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The precipitation in {@ds1@} Eastern African cities in August was measured and the results were grouped in intervals of 6mm. The results are displayed in the graph below. Place the mode line in the modal value 'by eye;.</p>

<p>
[[jsxgraph  input-ref-ans1="sR1" input-ref-ans2="sR2"]] 
JXG.Options.text.useMathJax = true; 

// Create empty list of points
var p = [];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [{#xmin#}, {#ymax#}, {#xmax#}, {#ymin#}],
  keepaspectratio: false,
  showCopyright: false,
  axis:true,
  pan:{enabled:false}
});

// Create points for the Data Set
var lx1 = {#dorx1#};
var ly1 = {#dory#};
var i = 0;
for (i = 0; i <= lx1.length - 1; i++) {
  p[i] = board.create('point', [lx1[i], ly1[i]], {
    face: 'x',
    name: '',
    size: 2,
    fixed:true
  });
}

// Create hidden point for the mean
var qm1 = board.create('point', [5,0], {
  name: "C",
  size: 2,
  visible: false
});

// Line for the mean
var qm11 = board.create('point', [5,5], {
  name: "C",
  size: 2,
  visible: false
});
var m1 =  board.create('line', [qm1, qm11], {
  straightFirst: true,
  straightLast: true,
  strokeWidth: 2,
  strokeColor: '#21a34c'
});
var lblm1 = board.create('text', [function() {return qm1.X()+1},{#ymax-1#},'mode'], { 
  fontSize: 15,
  color: '#21a34c'
});

/* Set answers from state of graph */
stack_jxg.bind_point(sR1,qm1);
stack_jxg.bind_point(sR2,qm11);
board.unsuspendUpdate();


[[/jsxgraph]]
</p>

<p style="display:none;">[[input:ans1]] [[validation:ans1]] </p>
<p style="display:none;">[[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Parameters for the distribution */
mean1:rand(30)+30;
sd1:(rand(0.3)+0.5)*mean1;
ds1:rand(50)+50;
bb:(sd1^2/mean1);
aa:(mean1^2/sd1^2);

/* Create the data using with a left skew */
dunx:makelist(round(float(quantile_gamma(rand(1.0),aa,bb))),i,1,ds1);
dorx:sort(dunx);

/* Group data on intervals of length 6 */
dorx1:makelist(1,i,1,ds1);
for i:1 thru ds1 do dorx1[i]:6*floor(dorx[i]/6)+3;

/* Calculate the actual mean of the grouped data */
m1:float(mean(dorx1));
mr1:round(m1)

/* Calculate the actual median of the grouped data */
m2:float(median(dorx1));
mr2:round(m2)

/* Finding bounds for the graph */
/* Calculating the maximum frequency to determine the max y-value */
maxfreq1:0;
k:1;
k_x:0;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (k:k+1, if maxfreq1 < k then (maxfreq1:k, k_x:dorx1[i])) else k:1;
/* Define max y-value. Set it to a minimum of 20 so that points are not squashed together */
ymax:if maxfreq1+2<20 then 20 else maxfreq1+2;
/* Define the min y-value and the max and min x-values */
xmax:dorx1[ds1]+2;/*+(dorx[ds1]-dorx[1])/1.5 can be added to create empty space on the right to make it visible on mobile phones */
xmin:if dorx1[1]-10<0 then dorx1[1]-10 else -10;
ymin:-2;

/* Find the y-coordinate for each point */
j:1;
dory:makelist(0,i,1,ds1);
dory[1]:1;
for i:2 thru ds1 do if dorx1[i]=dorx1[i-1] then (j:j+1, dory[i]:j) else (j:1, dory[i]:j);

/*finding the mode*/
m3:mode(dorx1);
m3:mean(m3);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Mean = "rand(30)+30" = {@mean1@}, SD = "(rand(0.3)+0.5)*Mean" = {@sd1@}</p> 
<p>Data Size = "rand(50)+50" = {@ds1@}</p>
<p>Gamma Parameters [a,b] = [{@aa@},{@bb@}]</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>m3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>maxfreq</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>mult_max: if length(mode(dorx1))=1 then 0 else 1;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumRelative</answertest>
        <sans>ans1[1]</sans>
        <tans>m3</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer for the mode is close enough. The actual value is {@m3@} mm.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>NumRelative</answertest>
        <sans>ans1[1]</sans>
        <tans>m1</tans>
        <testoptions>0.2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Your answer for the mode is almost close enough. The actual value is {@m3@} mm, with practice you should be able to estimate the mode to within 10%.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>mult_max</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>If there is more than one modal value then you should take the mean of the modal values as the mode.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>You are a bit too far from the actual mode value of {@mr1@} mm. Your estimate of the mode is more than 20% away.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1891023912</deployedseed>
    <deployedseed>1036719287</deployedseed>
    <deployedseed>962584701</deployedseed>
    <deployedseed>1748031846</deployedseed>
    <deployedseed>1773755093</deployedseed>
    <deployedseed>561483155</deployedseed>
    <deployedseed>1985349951</deployedseed>
    <deployedseed>1709020466</deployedseed>
    <deployedseed>1977300406</deployedseed>
    <deployedseed>105565799</deployedseed>
    <deployedseed>1639130340</deployedseed>
    <deployedseed>510985240</deployedseed>
    <deployedseed>961416529</deployedseed>
    <deployedseed>1704670646</deployedseed>
    <deployedseed>651592354</deployedseed>
    <deployedseed>758171530</deployedseed>
    <deployedseed>524304928</deployedseed>
    <deployedseed>1011660361</deployedseed>
    <deployedseed>133693342</deployedseed>
    <deployedseed>735813825</deployedseed>
    <deployedseed>232826220</deployedseed>
    <deployedseed>1890592215</deployedseed>
    <deployedseed>484655012</deployedseed>
    <deployedseed>2040736930</deployedseed>
    <deployedseed>673669126</deployedseed>
  </question>

</quiz>